# Week 9: Data Visualisation using Matplotlib

## Introduction

Data visualisation is a key aspect of mathematical computing and scientific programming.  
In this assignment, you will:
- Plot mathematical functions using `matplotlib`
- Create scatter plots
- Add labels, titles, legends, and annotations
- Visualise mathematical sequences

## Instructions

1. Do not change the function names.
2. Do not add or remove parameters from function definitions.
3. Plots should be generated using `matplotlib.pyplot`.
4. Each function must return the figure object for testing purposes.
5. Run all cells to verify outputs.

## 1: Plot Sine Function (5 marks)

Write a function `plot_sine` that:
- Takes no arguments
- Plots the function `y = sin(x)` for `x` from `0` to `2π`
- Labels the x-axis as `"x"`, y-axis as `"sin(x)"`, and title as `"Sine Function"`
- Returns the figure object

```python
import numpy as np
import matplotlib.pyplot as plt

def plot_sine():
    
```

<!-- grade_id: cell-b1772873e629ed11, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
fig = plot_sine()
assert fig is not None
plt.close(fig)
### END HIDDEN TESTS
```

## 2: Exponential Plot (5 marks)

Write a function `plot_exponential` that:
- Plots the function `y = e^x` for `x` in the range `0` to `5`
- Titles the plot `"Exponential Function"`
- Adds labels `"x"` and `"exp(x)"` on respective axes
- Returns the figure object

```python
def plot_exponential():
```

<!-- grade_id: cell-6c4733386e6b6416, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
fig = plot_exponential()
assert fig is not None
plt.close(fig)
### END HIDDEN TESTS
```

## 3: Scatter Plot of Random Points (5 marks)

Write a function `scatter_points` that:
- Takes two lists `x_vals` and `y_vals`
- Creates a scatter plot
- Adds the title `"Scatter Plot"`
- Labels the axes appropriately
- Returns the figure object


```python
def scatter_points(x_vals, y_vals):
```

<!-- grade_id: cell-0dda63630e45fa8f, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
fig = scatter_points([1,2], [3,4])
assert fig is not None
plt.close(fig)
### END HIDDEN TESTS
```

## 4: Annotate a Point on a Graph (5 marks)

Write a function `annotate_graph` that:
- Plots `y = x^2` for `x` in range -5 to 5
- Annotates the point (2, 4) with the label `"Point (2,4)"`
- Returns the figure object

```python
def annotate_graph():
```

<!-- grade_id: cell-bce0ee724f5395dd, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
fig = annotate_graph()
assert fig is not None
plt.close(fig)
### END HIDDEN TESTS
```

## 5: Visualise a Mathematical Sequence (5 marks)

Write a function `plot_sequence` that:
- Takes an integer `n`
- Generates the first `n` terms of the sequence: `a_n = n^2 + 1`
- Plots it as a line graph with markers
- Labels the sequence as `"a_n = n^2 + 1"`
- Returns the figure object

### Example
For n = 5 → The sequence is [2, 5, 10, 17, 26]

```python
def plot_sequence(n: int):
```

<!-- grade_id: cell-60e30ae51473dfae, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
fig = plot_sequence(5)
assert fig is not None
plt.close(fig)
### END HIDDEN TESTS
```