# Week 5: Functions

## Introduction

In this assignment, you will practice writing and using **functions** in Python.  
Functions are reusable blocks of code that help in breaking down problems into smaller parts.  

By the end of this lab, you will learn to:
- Define and call functions
- Pass parameters to functions
- Return values from functions
- Write functions for mathematical sequences
- Generate prime numbers using functions

This lab is designed to be completed in **2 hours**. Please read the instructions carefully and attempt all exercises.

## Instructions

1. Each coding problem must be solved inside the provided code cells.
2. Do not change the function names provided.
3. Do not add or remove arguments from the given function signatures.
4. Run all cells to check your progress.
5. Your code will be **autograded**, so follow the instructions strictly.

## Problem 1: Basic Function Definition (10 marks)

Write a function called `greet` that takes a string argument `name` and returns a greeting message in the form:

```
Hello, <name>!
```

### Example
```python
greet("Alice") → "Hello, Alice!"
```

```python
def greet():
    pass
```

<!-- grade_id: cell-7bb232a3e9809c4c, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert greet("Alice") == "Hello, Alice!"
assert greet("Bob") == "Hello, Bob!"
### END HIDDEN TESTS
```

## Problem 2: Function with Parameters (10 marks)

Write a function called `power` that takes two integers `a` and `b`, and returns `a` raised to the power `b`.

### Example
```python
power(2, 3) → 8
power(5, 0) → 1
```

```python
def power():
    pass
```

<!-- grade_id: cell-1ea57b390bf25845, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert power(2, 3) == 8
assert power(5, 0) == 1
assert power(3, 4) == 81
### END HIDDEN TESTS
```

## Problem 3: Factorial Function (10 marks)

Write a function `factorial` that takes a non-negative integer `n` and returns the factorial of `n`.

**Recall:**  
n! = n × (n-1) × (n-2) × ... × 1
and `0! = 1`.

### Example
```python
factorial(5) → 120
```

```python
def factorial():
    pass
```

<!-- grade_id: cell-ff7e643424cc6635, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert factorial(0) == 1
assert factorial(1) == 1
assert factorial(5) == 120
assert factorial(6) == 720
### END HIDDEN TESTS
```

## Problem 4: Fibonacci Sequence (10 marks)

Write a function `fibonacci` that takes an integer `n` and returns the `n`-th Fibonacci number.

**Recall:**  
The Fibonacci sequence is defined as:
F(0) = 0, F(1) = 1
F(n) = F(n-1) + F(n-2) for n ≥ 2

### Example
```python
fibonacci(0) → 0
fibonacci(1) → 1
fibonacci(6) → 8
```

```python
def fibonacci():
    pass
```

<!-- grade_id: cell-0cc30c2a0d4193a5, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(6) == 8
assert fibonacci(10) == 55
### END HIDDEN TESTS
```

## Problem 5: Prime Number Check (10 marks)

Write a function `is_prime` that takes an integer `n` and returns `True` if `n` is a prime number, and `False` otherwise.

**Recall:**  
A prime number is a natural number greater than 1 that is not divisible by any number other than 1 and itself.

### Example
```python
is_prime(7) → True
is_prime(10) → False
```

```python
def is_prime():
    pass
```

<!-- grade_id: cell-735e035bfb0176cb, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert is_prime(2) == True
assert is_prime(3) == True
assert is_prime(4) == False
assert is_prime(29) == True
### END HIDDEN TESTS
```

## Problem 6: Generate Prime Numbers (10 marks)

Write a function `generate_primes` that takes an integer `n` and returns a list of the first `n` prime numbers.

### Example
```python
generate_primes(5) → [2, 3, 5, 7, 11]
```

```python
def generate_primes():
    pass
```

<!-- grade_id: cell-2e27ec292006d68e, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert generate_primes(1) == [2]
assert generate_primes(3) == [2, 3, 5]
assert generate_primes(5) == [2, 3, 5, 7, 11]
assert generate_primes(7) == [2, 3, 5, 7, 11, 13, 17]
### END HIDDEN TESTS
```