# Week 2: Variables, Expressions, and Statements

Welcome to **Week 2** of the Python Lab for MSc Mathematics students.  

In this lab, you will:

- Learn how to use variables in Python.  
- Practice writing arithmetic expressions.  
- Work with Python statements.  
- Solve problems involving geometric shapes.  

This lab is designed to be completed in **1 hour**. Please read the instructions carefully and attempt all exercises.  

---

## Learning Objectives

By the end of this lab, you should be able to:

1. Assign values to variables and use them in expressions.  
2. Perform simple arithmetic operations in Python.  
3. Write Python statements to compute areas and perimeters.  
4. Develop problem-solving skills using variables and expressions.  

---

## Guidelines for Submissions

- Write your code **only** in the space provided in code cells.  
- Do not modify the test cells.  
- Run all cells before submission to ensure correctness.  
- Partial credit may be awarded if your code is logically correct but incomplete.  

---

# Part A: Variables and Arithmetic Expressions

In this section, we will practice simple variable assignments and arithmetic.


```python
# Q1. Assigning and Printing Variables (5 marks)
# Create three variables x, y, and z with the values 5, 10, and 15.
# Then print them in a single line, separated by spaces.

# YOUR CODE HERE
x = 5
y = 10
z = 15
print(x, y, z)
```

Output:
```
5 10 15
```

<!-- grade_id: cell-f9c1a88999e10311, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert x == 5
assert y == 10
assert z == 15
### END HIDDEN TESTS
```

```python
# Q2. Simple Arithmetic (5 marks)
# Compute result = (a + b) * (a - b) where a = 12 and b = 8.

# YOUR CODE HERE
a = 12
b = 8
result = (a + b) * (a - b)
result
```

Output:
```
80
```

<!-- grade_id: cell-3d55c5409a4c3b97, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert result == (12 + 8) * (12 - 8)
### END HIDDEN TESTS
```

```python
# Q3. Division and Remainder (5 marks)
# For m = 29 and n = 5, compute quotient and remainder.

# YOUR CODE HERE
m = 29
n = 5
quotient = m // n
remainder = m % n
(quotient, remainder)
```

Output:
```
(5, 4)
```

<!-- grade_id: cell-ff5dc9a618d915f1, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert quotient == 29 // 5
assert remainder == 29 % 5
### END HIDDEN TESTS
```

# Part B: Geometry with Variables

We will now use Python to compute areas and perimeters of geometric shapes.

```python
# Q4. Area and Perimeter of a Rectangle (5 marks)
# length = 12, width = 8

# YOUR CODE HERE
length = 12
width = 8
area_rectangle = length * width
perimeter_rectangle = 2 * (length + width)
(area_rectangle, perimeter_rectangle)
```

Output:
```
(96, 40)
```

<!-- grade_id: cell-8a56ce34defe9567, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert area_rectangle == 96
assert perimeter_rectangle == 40
### END HIDDEN TESTS
```

```python
# Q5. Area and Circumference of a Circle (5 marks)
# radius r = 7

import math

# YOUR CODE HERE
r = 7
area_circle = math.pi * r**2
circumference_circle = 2 * math.pi * r
(area_circle, circumference_circle)
```

Output:
```
(153.93804002589985, 43.982297150257104)
```

<!-- grade_id: cell-29ebb4342da62834, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert round(area_circle, 2) == round(math.pi * 7**2, 2)
assert round(circumference_circle, 2) == round(2 * math.pi * 7, 2)
### END HIDDEN TESTS
```

```python
# Q6. Area of a Triangle (Heron’s Formula) (5 marks)
# Sides: a = 13, b = 14, c = 15

import math

# YOUR CODE HERE
a, b, c = 13, 14, 15
s = (a + b + c) / 2
area_triangle = math.sqrt(s * (s - a) * (s - b) * (s - c))
area_triangle
```

Output:
```
84.0
```

<!-- grade_id: cell-81d9361745e35217, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
s_test = (13+14+15)/2
assert round(area_triangle, 2) == round(math.sqrt(s_test*(s_test-13)*(s_test-14)*(s_test-15)), 2)
### END HIDDEN TESTS
```

# Part C: Practice Problems

Apply your knowledge of variables and expressions to solve slightly more challenging problems.

```python
# Q7. Temperature Conversion
# Convert 37°C to Fahrenheit. (5 marks)

# YOUR CODE HERE
C = 37
fahrenheit = (C * 9/5) + 32
fahrenheit
```

Output:
```
98.6
```

<!-- grade_id: cell-2c97727672eaac13, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert fahrenheit == (9/5)*37 + 32
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[14], line 2
      1 ### BEGIN HIDDEN TESTS
----> 2 assert fahrenheit == (9/5)*37 + 32
      3 ### END HIDDEN TESTS

AssertionError:
```

```python
# Q8. Compound Interest
# P = 1000, r = 5%, t = 2 years (5 marks)

# YOUR CODE HERE
P = 1000
r = 5
t = 2
amount = P * (1 + r/100)**t
amount
```

Output:
```
1102.5
```

<!-- grade_id: cell-b0fb25f1d246efb8, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert round(amount, 2) == round(1000 * (1 + 0.05)**2, 2)
### END HIDDEN TESTS
```

```python
# Q9. Speed, Distance, and Time
# Distance = 150 km, Time = 3 hours (5 marks)

# YOUR CODE HERE
distance = 150
time = 3
speed = distance / time
speed
```

Output:
```
50.0
```

<!-- grade_id: cell-c186ca1812cd28b6, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert speed == 50
### END HIDDEN TESTS
```