# Week 3: Conditionals and Boolean Expressions

Welcome to Week 3 of the Python Programming Lab for MSc Mathematics students.  
In this lab, we will practice **Boolean expressions** and **conditional statements (`if`, `elif`, `else`)**.  

By the end of this lab, you should be able to:
- Understand Boolean values (`True`, `False`) and relational operators.  
- Write conditional expressions using `if`, `elif`, and `else`.  
- Solve problems like determining whether a number is even/odd.  
- Use conditionals to classify triangles.  

---

## Part 1: Boolean Expressions

Boolean expressions evaluate to `True` or `False`.  
Examples:
- `5 > 3` → `True`  
- `10 == 2*5` → `True`  
- `7 < 2` → `False`  

Try running the following examples:

```python
# Example: Boolean expressions
print(5 > 3)
print(10 == 2*5)
print(7 < 2)
```

Output:
```
True
True
False
```

## Part 2: Using `if`, `elif`, and `else`

The `if` statement allows your program to make decisions.

**Syntax:**

```python
if (condition):
    # block executed if condition is True
elif (another_condition):
    # block executed if previous condition was False but this is True
else:
    # block executed if none of the above are True
```

```python
# Example: Checking if a number is positive, negative, or zero
x = 5
if x > 0:
    print("Positive")
elif x == 0:
    print("Zero")
else:
    print("Negative")
```

Output:
```
Positive
```

## Part 3: Assignments

### Q1. Even or Odd (10 marks)

Write a function `even_or_odd(n)` that returns:
- `"Even"` if `n` is even  
- `"Odd"` if `n` is odd  

**Example:**
- `even_or_odd(4)` → `"Even"`
- `even_or_odd(7)` → `"Odd"`

```python
def even_or_odd(n: int) -> str:
    """
    Return "Even" if the number is even, otherwise return "Odd".
    """
    # YOUR CODE HERE
    if (n % 2) == 0:
        return "Even"
    else:
        return "Odd"

# Test cases
print(even_or_odd(2))  # Expected "Even"
print(even_or_odd(7))  # Expected "Odd"
```

Output:
```
Even
Odd
```

<!-- grade_id: cell-f3fc8e60df1af6c5, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert even_or_odd(10) == "Even"
assert even_or_odd(13) == "Odd"
assert even_or_odd(0) == "Even"
### END HIDDEN TESTS
```

### Q2. Maximum of Three Numbers (10 marks)

Write a function `max_of_three(a, b, c)` that returns the largest among `a`, `b`, and `c`.

**Example:**
- `max_of_three(3, 7, 5)` → `7`

```python
def max_of_three(a: int, b: int, c: int) -> int:
    """
    Return the largest among a, b, and c.
    """
    if a >= b and a >= c:
        return a
    elif b >= a and b >= c:
        return b
    else:
        return c

# Test cases
print(max_of_three(3, 7, 5))    # Expected 7
print(max_of_three(-1, -5, -3)) # Expected -1

```

Output:
```
7
-1
```

<!-- grade_id: cell-5fd75c56aaaa166d, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert max_of_three(10, 20, 30) == 30
assert max_of_three(50, 20, 10) == 50
assert max_of_three(7, 7, 7) == 7
### END HIDDEN TESTS
```

### Q3. Leap Year Checker (10 marks)

Write a function `is_leap_year(year)` that returns:
- `True` if the year is a leap year  
- `False` otherwise  

**Rules for leap year:**
- Year divisible by 4 → leap year  
- But if divisible by 100 → not a leap year  
- Unless also divisible by 400 → leap year  

**Example:**
- `is_leap_year(2000)` → `True`  
- `is_leap_year(1900)` → `False`  
- `is_leap_year(2024)` → `True`

```python
def is_leap_year(year: int) -> bool:
    """
    Return True if year is a leap year, else False.
    
    A leap year is:
    - Divisible by 4
    - Not divisible by 100 unless also divisible by 400
    """
    if year % 4 != 0:
        return False
    elif year % 100 != 0:
        return True
    elif year % 400 == 0:
        return True
    else:
        return False

# Test cases 
print(is_leap_year(2000))  # Expected True
print(is_leap_year(1900))  # Expected False
print(is_leap_year(2024))  # Expected True

```

Output:
```
True
False
True
```

<!-- grade_id: cell-9d37ae1ad21637b1, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert is_leap_year(2100) == False
assert is_leap_year(2400) == True
assert is_leap_year(2023) == False
### END HIDDEN TESTS
```

### Q4. Triangle Classifier (20 marks)

Write a function `classify_triangle(a, b, c)` that classifies a triangle based on side lengths:
- `"Equilateral"` → all sides equal  
- `"Isosceles"` → exactly two sides equal  
- `"Scalene"` → all sides different  
- `"Not a Triangle"` → if the given sides do not form a valid triangle  

**Hint:** For three sides `a, b, c` to form a triangle:
- `a + b > c`, `a + c > b`, and `b + c > a` must all hold true.

**Examples:**
- `classify_triangle(3, 3, 3)` → `"Equilateral"`  
- `classify_triangle(3, 4, 4)` → `"Isosceles"`  
- `classify_triangle(3, 4, 5)` → `"Scalene"`  
- `classify_triangle(1, 2, 3)` → `"Not a Triangle"`

```python
def classify_triangle(a: int, b: int, c: int) -> str:
    """
    Classify the triangle based on side lengths.
    Returns:
        - "Equilateral" if all sides are equal
        - "Isosceles" if exactly two sides are equal
        - "Scalene" if all sides are different
        - "Not a Triangle" if the sides do not form a valid triangle
    """
    # Check for triangle inequality
    if a + b <= c or a + c <= b or b + c <= a:
        return "Not a Triangle"
    
    if a == b == c:
        return "Equilateral"
    elif a == b or b == c or a == c:
        return "Isosceles"
    else:
        return "Scalene"

# Test cases
print(classify_triangle(3, 3, 3))  # Expected "Equilateral"
print(classify_triangle(3, 4, 4))  # Expected "Isosceles"
print(classify_triangle(1, 2, 3))  # Expected "Not a Triangle"

```

Output:
```
Equilateral
Isosceles
Not a Triangle
```

<!-- grade_id: cell-ab41c168aee1f669, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert classify_triangle(3, 4, 5) == "Scalene"
assert classify_triangle(10, 10, 15) == "Isosceles"
assert classify_triangle(5, 5, 5) == "Equilateral"
assert classify_triangle(2, 2, 5) == "Not a Triangle"
### END HIDDEN TESTS
```

## Part 4: Reflection Questions (No Grading)

Answer the following in your own words (write in a markdown cell):

1. What is the difference between `if` and `elif`?  
2. Why do we need the `else` statement? Can we write conditionals without it?  
3. What checks must be satisfied for three numbers to form a triangle?  