# Week 1 Assignment: Python Installation and First Steps

## Maximum Marks - 25 points
## Duration - 1 hour

Welcome to your first programming assignment in Python.

This notebook will guide you through some basic steps:  
- Checking your Python installation  
- Writing and running your first program  
- Using Jupyter notebooks effectively  
- Exploring the built-in `help()` function  

Each problem below has instructions in **markdown** and space for you to write your **Python code**.  
Make sure you follow the instructions carefully.  

## Problem 1: Check Python Version (5 points)

Before we start coding, let's make sure Python is installed correctly.  

**Task:**  
Write code to print the version of Python that is currently running.  
Hint: You can use the `sys` module.

```python
# Problem 1

import sys

# YOUR CODE HERE
print(sys.version)  ## Already written, just run this cell.
```

<!-- grade_id: cell-9bb6a273b213ff89, points: 5, grade: true, solution: false -->
```python

```

## Problem 2: Write Your First Program (10 points)

Let’s write your very first Python program.  

**Task:**  
Write a function called `greet()` that returns the string:  

```text
Hello, MSc Mathematics Students!
```

*Do not print inside the function, just return the string.*  

```python
# Problem 2

def greet():
    # YOUR CODE HERE
    return "<put your string here>"  # Replace with the correct string
```

<!-- grade_id: cell-dd7e347cf3bd1f50, points: 10, grade: true, solution: false -->
```python
# DO NOT EDIT THIS CELL
# TEST CASE for Problem 2
assert greet() == "Hello, MSc Mathematics Students!"
```

## Problem 3: Print a Custom Greeting (10 points)

Functions often take input from the user.  

**Task:**  
Write a function called `personal_greet(name)` that takes a string `name`  
and returns `"Hello, <name>!"`  

Example:  
- `personal_greet("Alice")` → `"Hello, Alice!"`

```python
# Problem 3

def personal_greet(name):
    # YOUR CODE HERE
    return f"<put the correct string here>, {name}!" # Replace with the correct string
```

<!-- grade_id: cell-abffe08c70dc4a5f, points: 10, grade: true, solution: false -->
```python
# DO NOT EDIT THIS CELL. 
# TEST CASE for Problem 3
assert personal_greet("Alice") == "Hello, Alice!"
assert personal_greet("Bob") == "Hello, Bob!"
```

## Tutorial: Understanding `help()` Command

Python has a built-in `help()` function that lets you explore documentation. 

For usage, just call the `help()` function in new cell below, you will be prompted to enter the module name. Write any module name of your choice like `Math` or `numpy`. If you want to explore what are the modules present, write `modules`. Explore the `help()` function below and use it in day-to-day practice to understand the concepts better with the help of official Python documentation. 

*Note - This is not graded.*